# Assignment 2

For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

```python
import pandas as pd

df = pd.read_csv('assets/NISPUF17.csv')
df['EDUC1']

full_check = df['EDUC1'].notna().all()
if(full_check):
    print("Vsetky hodnoty su vyplnene")
    
missing_values = df['EDUC1'].isna().sum()
missing_values

print(len(df['EDUC1']))
df['EDUC1']
```

Output:
```
Vsetky hodnoty su vyplnene
28465
0        4
1        3
2        3
3        4
4        1
        ..
28460    3
28461    2
28462    3
28463    2
28464    4
Name: EDUC1, Length: 28465, dtype: int64
```

```python
import pandas as pd

df = pd.read_csv('assets/NISPUF17.csv')

less_than_high_school = 0
high_school = 0
more_than_high_school_but_not_college = 0
college = 0
length = len(df['EDUC1'])

for value in df['EDUC1']:
    if value == 1:
        less_than_high_school += 1
    elif value == 2:
        high_school += 1
    elif value == 3:
        more_than_high_school_but_not_college += 1
    elif value == 4:
        college += 1

result = {"less than high school":less_than_high_school / length,
        "high school":high_school / length,
        "more than high school but not college":more_than_high_school_but_not_college / length,
        "college":college / length}
    
print(result)
```

Output:
```
{'less than high school': 0.10202002459160373, 'high school': 0.172352011241876, 'more than high school but not college': 0.24588090637625154, 'college': 0.47974705779026877}
```

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```

<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
def proportion_of_education():
    less_than_high_school = 0
    high_school = 0
    more_than_high_school_but_not_college = 0
    college = 0
    
    length = len(df['EDUC1'])
    
    for value in df['EDUC1']:
        if value == 1:
            less_than_high_school += 1
        elif value == 2:
            high_school += 1
        elif value == 3:
            more_than_high_school_but_not_college += 1
        elif value == 4:
            college += 1
            
    result = {"less than high school":less_than_high_school / length,
        "high school":high_school / length,
        "more than high school but not college":more_than_high_school_but_not_college / length,
        "college":college / length}
            
    return result
    # your code goes here
    # YOUR CODE HERE
    # raise NotImplementedError()
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

```python
# df[['CBF_O1', 'INFLUENZA']]
# columns_starting_with_p_num = [col for col in df.columns if col.startswith('P_NUM')]

# v tomto by mali byt vsetky tie hodnoty, ktore potrebujem, cbf je ci boli kojeni a ostatne s tou jednou vakcinou, ale rozne typy danej vakciny
looking_for = [col for col in df.columns if col.startswith('P_NUMFL')]
looking_for.insert(0, 'CBF_01')
print(looking_for)

df[looking_for]

rozne_hodnoty = set()
count77 = 0
count99 = 0

# tu vobec netusim co robi 77 a 99
for value in df['CBF_01']:
    rozne_hodnoty.add(value)
    if value == 77:
        count77 += 1
    elif value == 99:
        count99 += 1
    else:
        pass
    
print(f"count77 {count77}")
print(f"count99 {count99}")

print(rozne_hodnoty)
# df['FL']
```

Output:
```
['CBF_01', 'P_NUMFLU', 'P_NUMFLUL', 'P_NUMFLUM', 'P_NUMFLUN']
count77 110
count99 9
{1, 2, 99, 77}
```

```python
import numpy as np

rozne_hodnoty2 = set()
counter = 0

# tu vobec netusim co robi 77 a 99
for value in df['P_NUMFLU']:
    rozne_hodnoty2.add(value)
    if not np.isnan(value):
        counter += 1

print(counter)
# P_NUMFLU by malo byt total count of every type tej vakciny, ktoru hladame
```

Output:
```
15333
```

```python
import numpy as np

for value in df[['CBF_01', 'P_NUMFLU']]:
    print(value)

cbf_counter_1 = 0
cbf1_total_value_count = 0
cbf_counter_2 = 0
cbf2_total_value_count = 0

for index, row in df.iterrows():
    if row['CBF_01'] == 1 and not np.isnan(row['P_NUMFLU']):
        cbf_counter_1 += 1
        cbf1_total_value_count += row['P_NUMFLU']
    elif row['CBF_01'] == 2 and not np.isnan(row['P_NUMFLU']):
        cbf_counter_2 += 1
        cbf2_total_value_count += row['P_NUMFLU']
        
print(cbf_counter_1, cbf_counter_2) 
print(cbf1_total_value_count, cbf2_total_value_count)
breast_milk_vaccines_average = cbf1_total_value_count / cbf_counter_1
no_breast_milk_vaccines_average = cbf2_total_value_count / cbf_counter_2

print(breast_milk_vaccines_average, no_breast_milk_vaccines_average)
```

Output:
```
CBF_01
P_NUMFLU
13291 1997
24986.0 3188.0
1.8799187420058687 1.5963945918878317
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
import numpy as np

def average_influenza_doses():
    cbf_counter_1 = 0
    cbf1_total_value_count = 0
    cbf_counter_2 = 0
    cbf2_total_value_count = 0

    for index, row in df.iterrows():
        if row['CBF_01'] == 1 and not np.isnan(row['P_NUMFLU']):
            cbf_counter_1 += 1
            cbf1_total_value_count += row['P_NUMFLU']
        elif row['CBF_01'] == 2 and not np.isnan(row['P_NUMFLU']):
            cbf_counter_2 += 1
            cbf2_total_value_count += row['P_NUMFLU']

    # print(cbf_counter_1, cbf_counter_2) 
    # print(cbf1_total_value_count, cbf2_total_value_count)
    breast_milk_vaccines_average = cbf1_total_value_count / cbf_counter_1
    no_breast_milk_vaccines_average = cbf2_total_value_count / cbf_counter_2
    
    return (breast_milk_vaccines_average, no_breast_milk_vaccines_average)
    # raise NotImplementedError()
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

```python
import pandas as pd

df = pd.read_csv('assets/NISPUF17.csv')

rozne_hodnoty = set()

for value in df['HAD_CPOX']:
    rozne_hodnoty.add(value)
    
print(rozne_hodnoty)
# df['HAD_CPOX']

df['P_NUMVRC'].dropna()
```

Output:
```
{1, 2, 99, 77}
3        1.0
4        0.0
6        1.0
7        0.0
8        0.0
        ... 
28446    1.0
28447    1.0
28449    1.0
28452    1.0
28454    1.0
Name: P_NUMVRC, Length: 15333, dtype: float64
```

```python
had_chickenpox_male = 0
had_chickenpox_female = 0
didnt_have_chickenpox_male = 0
didnt_have_chickenpox_female = 0

for index, row in df.iterrows():
    if row['P_NUMVRC'] >= 1:
        if row['HAD_CPOX'] == 1:
            if row['SEX'] == 1:
                had_chickenpox_male += 1
            elif row['SEX'] == 2:
                had_chickenpox_female += 1
        elif row['HAD_CPOX'] == 2:
            if row['SEX'] == 1:
                didnt_have_chickenpox_male += 1
            elif row['SEX'] == 2:
                didnt_have_chickenpox_female += 1

# print(had_chickenpox)
# print(didnt_have_chickenpox)
print(had_chickenpox_male, had_chickenpox_female)
print(didnt_have_chickenpox_male, didnt_have_chickenpox_female)  

# 1 MALE, 2 FEMALE
df['SEX'].head()

# NAD TYMTO POROZMYSLAT
# VECTOROVE OPERACIE SU VZDY LEPSIE V PANDAS

# Filtrovanie iba očkovaných detí (aspoň 1 dávka vakcíny)
vaccinated = df[df["P_NUMVRC"] >= 1]

# Počet očkovaných detí podľa pohlavia a toho, či mali ovčie kiahne
counts = vaccinated.groupby(["SEX", "HAD_CPOX"]).size().unstack(fill_value=0)

# Získanie hodnôt podľa pohlavia
had_chickenpox_male = counts.loc[1, 1] if 1 in counts.index else 0
had_chickenpox_female = counts.loc[2, 1] if 2 in counts.index else 0
didnt_have_chickenpox_male = counts.loc[1, 2] if 1 in counts.index else 0
didnt_have_chickenpox_female = counts.loc[2, 2] if 2 in counts.index else 0

print(had_chickenpox_male, had_chickenpox_female)
print(didnt_have_chickenpox_male, didnt_have_chickenpox_female)  

print(had_chickenpox_male / didnt_have_chickenpox_male)
print(had_chickenpox_female / didnt_have_chickenpox_female)
```

Output:
```
68 53
7028 6802
68 53
7028 6802
0.009675583380762664
0.0077918259335489565
```

```python
vaccinated = df[df["P_NUMVRC"] >= 1]
# print(vaccinated)

counts = vaccinated.groupby(["SEX", "HAD_CPOX"]).size().unstack(fill_value=0)
print(counts)
```

Output:
```
HAD_CPOX  1     2   77
SEX                   
1         68  7028  22
2         53  6802  22
```

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex():
    # Filtrovanie iba očkovaných detí (aspoň 1 dávka vakcíny)
    vaccinated = df[df["P_NUMVRC"] >= 1]

    # Počet očkovaných detí podľa pohlavia a toho, či mali ovčie kiahne
    counts = vaccinated.groupby(["SEX", "HAD_CPOX"]).size().unstack(fill_value=0)

    # Získanie hodnôt podľa pohlavia
    had_chickenpox_male = counts.loc[1, 1] if 1 in counts.index else 0
    had_chickenpox_female = counts.loc[2, 1] if 2 in counts.index else 0
    didnt_have_chickenpox_male = counts.loc[1, 2] if 1 in counts.index else 0
    didnt_have_chickenpox_female = counts.loc[2, 2] if 2 in counts.index else 0

    result_male = had_chickenpox_male / didnt_have_chickenpox_male
    result_female = had_chickenpox_female / didnt_have_chickenpox_female
    
    result = {"male":result_male,
              "female":result_female}
    
    return result
    # raise NotImplementedError()
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more no’s) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses.

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isn’t really the full picture, since we are not looking at when the dose was given. It’s possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

```python
import scipy.stats as stats
import numpy as np
import pandas as pd

import pandas as pd

df = pd.read_csv('assets/NISPUF17.csv')

df = df.dropna(subset=["HAD_CPOX", "P_NUMVRC"])
df = df[df["HAD_CPOX"].isin([1, 2])]
print(len(df))
    
```

Output:
```
15286
```

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    df = pd.read_csv('assets/NISPUF17.csv')
    df = df.dropna(subset=["HAD_CPOX", "P_NUMVRC"])
    df = df[df["HAD_CPOX"].isin([1, 2])]
    
    # correlation
    corr, pval=stats.pearsonr(df["HAD_CPOX"],df["P_NUMVRC"])
    
    # just return the correlation
    return corr

    # YOUR CODE HERE
    # raise NotImplementedError()
```

```python
import scipy.stats as stats
import numpy as np
import pandas as pd
    
df = df.dropna(subset=["HAD_CPOX", "P_NUMVRC"])
df = df[df["HAD_CPOX"].isin([1, 2])]
    
# here is some stub code to actually run the correlation
corr, pval=stats.pearsonr(df["HAD_CPOX"],df["P_NUMVRC"])

print(corr)
```

Output:
```
0.07044873460147986
```

```python
df["HAD_CPOX"]
```

Output:
```
3        2
4        2
6        2
7        2
8        1
        ..
28446    2
28447    2
28449    2
28452    2
28454    2
Name: HAD_CPOX, Length: 15286, dtype: int64
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```

```python

```